# Assignment 2 - Network Connectivity

In this assignment you will go through the process of importing and analyzing an internal email communication network between employees of a mid-sized manufacturing company. 
Each node represents an employee and each directed edge between two nodes represents an individual email. The left node represents the sender and the right node represents the recipient. We will also store the timestamp of each email. 

```python
import networkx as nx

#!head assets/email_network.txt
```

### Question 1

Using networkx, load up the directed multigraph from `assets/email_network.txt`. Make sure the node names are strings.

*This function should return a directed multigraph networkx graph.*

<!-- grade_id: cell-60316968695c6f20, points: 0, grade: false, solution: true -->
```python
def answer_one():
    
    import pickle
    import matplotlib.pyplot as plt
    
    # YOUR CODE HERE
    
    G1 = nx.read_edgelist(path='assets/email_network.txt', data= [('time',int)],delimiter='\t', create_using=nx.MultiDiGraph())
    return G1
    raise NotImplementedError()
```

<!-- grade_id: cell-483a56c89f9db231, points: 1, grade: true, solution: false -->
```python
ans_one = answer_one()
```

### Question 2

How many employees are represented in the network?

How many `sender`->`recipient` pairs of employees are there in the network such that `sender` sent at least one email to `recipient`? Note that even if a `sender` sent multiple messages to a `recipient`, they should only be counted once. You should **not** exclude cases where an employee sent emails to themselves from this [email] count. 

*This function should return a tuple with two integers (#employees, # `sender`->`recipient` pairs).*

<!-- grade_id: cell-16a78f22207d9f95, points: 0, grade: false, solution: true -->
```python
def answer_two():
    # YOUR CODE HERE
    
    G2 = answer_one()
    lnode = G2.number_of_nodes()
    ledge = G2.number_of_edges()
    
    return lnode,ledge+nx.number_of_selfloops(G2)
    raise NotImplementedError()
```

<!-- grade_id: cell-5b6391549b076d2b, points: 1, grade: true, solution: false -->
```python
ans_two = answer_two()
ans_two
```

Output:
```
(167, 82978)
```

### Question 3

* Part 1. Assume that information in this company can only be exchanged through email.

    When an employee sends an email to another employee, a communication channel has been created, allowing the sender to provide information to the reciever, but not viceversa. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


* Part 2. Now assume that a communication channel established by an email allows information to be exchanged both ways. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


*This function should return a tuple of bools (part1, part2).*

<!-- grade_id: cell-606dd37fb2f55358, points: 0, grade: false, solution: true -->
```python
def answer_three():
    # YOUR CODE HERE
    
    A3 = answer_one()
    # P1
    # A -> B a directed edge created , one way directed 
    # strong_connected
    SC = bool(nx.is_strongly_connected(A3))

    # P2
    # weakly connected 
    WC = bool(nx.is_weakly_connected(A3))
    
    return (SC,WC)
    
    raise NotImplementedError()
```

<!-- grade_id: cell-82b3f0bc45e2895f, points: 1, grade: true, solution: false -->
```python
ans_three = answer_three()

```

### Question 4

How many nodes are in the largest weakly connected component of the graph?

*This function should return an int.*

<!-- grade_id: cell-39196f4c27b35534, points: 0, grade: false, solution: true -->
```python
def answer_four():
    # YOUR CODE HERE
    
    A4 = answer_one()
    largest_cc = max(nx.weakly_connected_components(A4), key=len)

    return len(largest_cc)
    raise NotImplementedError()
```

<!-- grade_id: cell-2b1b7b06ecfa751d, points: 1, grade: true, solution: false -->
```python
ans_four = answer_four()
```

### Question 5

How many nodes are in the largest strongly connected component?

*This function should return an int*

<!-- grade_id: cell-879d17a9db5c3790, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    A5 = answer_one()
    largest_scc = max(nx.strongly_connected_components(A5), key=len)

    return len(largest_scc)
    raise NotImplementedError()
```

<!-- grade_id: cell-b0524f7dc1fbdec4, points: 1, grade: true, solution: false -->
```python
ans_five = answer_five()

```

### Question 6

Using the NetworkX functions `strongly_connected_components` and `subgraph`, find the subgraph of nodes in the largest strongly connected component. 
Call this graph G_sc.

*This function should return a networkx MultiDiGraph named G_sc.*

<!-- grade_id: cell-a81e6f88ef84578c, points: 0, grade: false, solution: true -->
```python
def answer_six():
    # YOUR CODE HERE
    
    # getting the  largest SC node(s)
    # constructing the subgraph
    A6 = answer_one()
    lsc = max(nx.strongly_connected_components(A6), key=len)
    G_sc = nx.subgraph(A6,lsc)

    return G_sc
    raise NotImplementedError()


```

Output:
```
<function __main__.answer_six()>
```

<!-- grade_id: cell-cf148ef273b3b19c, points: 1, grade: true, solution: false -->
```python
ans_six = answer_six()
assert type(ans_six) == nx.MultiDiGraph , "Your return type should be a MultiDiGraph object"

```

### Question 7

What is the average distance between nodes in G_sc?

*This function should return a float.*

<!-- grade_id: cell-1a3075039b3c2c0b, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    #the strongly connected graph 
    A7 = answer_six()
    return nx.average_shortest_path_length(A7)
    raise NotImplementedError()
```

<!-- grade_id: cell-5b374fdd48f37e02, points: 1, grade: true, solution: false -->
```python
ans_seven = answer_seven()

```

### Question 8

What is the largest possible distance between two employees in G_sc?

*This function should return an int.*

<!-- grade_id: cell-4fbb8d34d00f675c, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    # YOUR CODE HERE
    A8 = answer_six()
    return nx.diameter(A8)
    raise NotImplementedError()
```

<!-- grade_id: cell-c5714787854ef644, points: 1, grade: true, solution: false -->
```python
ans_eight = answer_eight()

```

### Question 9

What is the set of nodes in G_sc with eccentricity equal to the diameter?

*This function should return a set of the node(s).*

<!-- grade_id: cell-bcfab850744277ed, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    A9 = answer_six()
    return set(nx.periphery(A9))
    raise NotImplementedError()
```

<!-- grade_id: cell-77c9ca0b94df3d6f, points: 1, grade: true, solution: false -->
```python
ans_nine = answer_nine()
assert type(ans_nine) == set, "Student answer must return a set"
ans_nine
```

Output:
```
{'129', '134', '97'}
```

### Question 10

What is the set of node(s) in G_sc with eccentricity equal to the radius?

*This function should return a set of the node(s).*

<!-- grade_id: cell-fd3d43ea37da03f7, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    A10 = answer_six()
    return set(nx.center(A10))
    raise NotImplementedError()
```

<!-- grade_id: cell-bfd2ee304bc25264, points: 1, grade: true, solution: false -->
```python
ans_ten = answer_ten()
assert type(ans_ten) == set, "Student answer must return a set"
```

### Question 11

Which node in G_sc has the most shortest paths to other nodes whose distance equal the diameter of G_sc?


For the node with the most such shortest paths, how many of these paths are there?


*This function should return a tuple (name of node, number of paths).*

<!-- grade_id: cell-c40858baf3d92139, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    # GsC = node eql to diameter
    
    # periphery is the set of nodes that have eccentricity ( largest distance between n and all other node ) to diameter
    A11 = answer_six()
    
    # req : find most num of SP for node(peri) and the num 
    
    # set of nodes
    per = nx.periphery(A11)
    pathNum ={}
    diam = nx.diameter(A11)
    
    # iterate through the set
    for node in per:
        #find SP of node(s)
        #networkx documentation
        # If only the source is specified, return a dictionary keyed by targets 
        # with a list of nodes in a shortest path from the source to one of the targets.
        short = nx.shortest_path(G=A11, source=node)
        # for every path generated in dict of short 
        # only accept path length == diameter [ 'by a shortest path of length eql to diameter ' ] 
        pathLeng = [ path for path in short.values() if (len(path) - 1 == diam ) ]
        # tag path length to the node key for referencing
        pathNum[node] = len(pathLeng)
    
    # dict key
    keys = list(pathNum.keys())
    values = list(pathNum.values())
    
    # identify maxkey -> keys number of max value
    maxKeyidentifier = keys[values.index(max(values))]
    
    return maxKeyidentifier, pathNum[maxKeyidentifier]
    raise NotImplementedError()
```

<!-- grade_id: cell-f79b06650f61cf37, points: 1, grade: true, solution: false -->
```python
ans_eleven = answer_eleven()
assert type(ans_eleven) == tuple, "Student answer must be a tuple"
```

### Question 12

Suppose you want to prevent communication from flowing to the node that you found in the previous question from any node in the center of G_sc, what is the smallest number of nodes you would need to remove from the graph (you're not allowed to remove the node from the previous question or the center nodes)?

*This function should return an integer.*

<!-- grade_id: cell-a4606db8971dcbba, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    
    #  using node connectivity 
    node = answer_eleven()[0]
    G = answer_one()
    
    connect = nx.node_connectivity(G, node,10)
    return connect
    raise NotImplementedError()
```

<!-- grade_id: cell-509cfa9f4136124d, points: 1, grade: true, solution: false -->
```python
ans_twelve = answer_twelve()
ans_twelve

```

Output:
```
File ~/Library/Python/3.9/lib/python/site-packages/networkx/classes/reportviews.py:421, in DiDegreeView.__init__(self, G, nbunch, weight)
    419 self._succ = G._succ if hasattr(G, "_succ") else G._adj
    420 self._pred = G._pred if hasattr(G, "_pred") else G._adj
--> 421 self._nodes = self._succ if nbunch is None else list(G.nbunch_iter(nbunch))
    422 self._weight = weight

File ~/Library/Python/3.9/lib/python/site-packages/networkx/classes/graph.py:2025, in Graph.nbunch_iter.<locals>.bunch_iter(nlist, adj)
   2021 if "hashable" in message:
   2022     exc = NetworkXError(
   2023         f"Node {n} in sequence nbunch is not a valid node."
   2024     )
-> 2025 raise exc

NetworkXError: nbunch is not a node or a sequence of nodes.
```

### Question 13

Convert the graph G_sc into an undirected graph by removing the direction of the edges of G_sc. Call the new graph G_un. 


*This function should return a networkx Graph.*

<!-- grade_id: cell-612b40425f0851e8, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    undirected = answer_six().to_undirected()
    return nx.Graph(undirected)
    raise NotImplementedError()
```

<!-- grade_id: cell-d1c0627a327cd774, points: 1, grade: true, solution: false -->
```python
ans_thirteen = answer_thirteen()
assert type(ans_thirteen) == nx.Graph , "Your return type should be a Graph object"

```

### Question 14

What is the transitivity and average clustering coefficient of graph G_un?

*This function should return a tuple (transitivity, avg clustering).*     
*Note: DO NOT round up your answer.*

<!-- grade_id: cell-63e48b1eed1a1fb8, points: 0, grade: false, solution: true -->
```python
def answer_fourteen():
    # YOUR CODE HERE
    A14 = answer_thirteen()
    return nx.transitivity(A14),nx.average_clustering(A14)
    raise NotImplementedError()
```

<!-- grade_id: cell-41dda91202f58e7e, points: 1, grade: true, solution: false -->
```python
ans_fourteen = answer_fourteen()
assert type(ans_fourteen) == tuple, "Student answer must be a tuple"

```